Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Gaston Itaka Kitambala"
COLLABORATORS = ""
```

---

###### <img style='margin-right:0' src="http://dinfo.ca/logoDptInfo.jpg" width=300>

# Laboratoire #10 (Python)
---

---
#### Question #1 (1 point)
Lors des Jeux du Québec 2018 à Thetford Mines, la finale des nageurs "Garçons - 13-15 ans, 200m libre" a homologué les temps suivants (exprimés sous forme de chaînes de caractères).

Vous devez trier l'ensemble de données et afficher uniquement les cinq meilleurs temps.

Résultats attendus:
```
Meilleurs temps
---------------
#1 -- 01:53.7
#2 -- 01:55.7
#3 -- 01:57.3
#4 -- 01:58.6
#5 -- 01:59.4
```

<!-- grade_id: q1-code, points: 0.5, grade: true, solution: true -->
```python
resultats = ['02:23.6', '02:25.7', '02:11.5', '02:09.6', '02:08.7', '01:55.7',
             '02:02.0', '02:20.9', '02:22.4', '02:04.2', '02:04.8', '02:18.5',
             '02:19.5', '02:12.3', '01:58.6', '01:53.7', '02:13.4', '02:06.9',
             '02:18.0', '02:16.1', '02:09.9', '02:07.1', '02:27.7', '01:57.3',
             '02:08.5', '02:18.7', '02:12.1', '02:04.2', '02:18.3', '02:48.7',
             '02:12.8', '02:16.3', '02:04.1', '02:25.6', '02:16.3', '02:37.2',
             '01:59.4']
print("Meilleurs temps")
print("-"*15)
resultats.sort()
for i in range(5):
    print("#%i -- %s" % (i+1, resultats[i]))
```

Output:
```
Meilleurs temps
---------------
#1 -- 01:53.7
#2 -- 01:55.7
#3 -- 01:57.3
#4 -- 01:58.6
#5 -- 01:59.4
```

<!-- grade_id: q1-tests, points: 0.5, grade: true, solution: false -->
```python
assert len(resultats) == 37, "Il devrait toujours y avoir 37 résultats"
assert resultats[0] == '01:53.7', 'Premier temps'
assert resultats[1] == '01:55.7', 'Deuxième temps'
assert resultats[2] == '01:57.3', 'Troisième temps'
assert resultats[3] == '01:58.6', 'Quatrième temps'
assert resultats[4] == '01:59.4', 'Cinquième temps'
assert resultats[36] == '02:48.7', 'Dernier temps'
```

---
#### Question #2 (1 point)

Pour la collection `bottinTelephonique`, on veut trier en ordre **décroissant** de numéro de téléphone.

Résultats attendus:

```['514-555-3333 Eric Fillion', '514-555-2222 Guillaume Huard', '418-555-1111 Kim Labrecque', '418-333-3333 Ian Jacob', '418-332-2222 Charles Demers', '418-331-1111 Adam Bernard']```

<!-- grade_id: q2-code, points: 0.5, grade: true, solution: true -->
```python
bottinTelephonique = ['418-332-2222 Charles Demers',
                      '514-555-3333 Eric Fillion',
                      '514-555-2222 Guillaume Huard',
                      '418-333-3333 Ian Jacob',
                      '418-555-1111 Kim Labrecque',
                      '418-331-1111 Adam Bernard']
bottinTelephonique.sort(reverse = True)
print(bottinTelephonique)
```

Output:
```
['514-555-3333 Eric Fillion', '514-555-2222 Guillaume Huard', '418-555-1111 Kim Labrecque', '418-333-3333 Ian Jacob', '418-332-2222 Charles Demers', '418-331-1111 Adam Bernard']
```

<!-- grade_id: q2-tests, points: 0.5, grade: true, solution: false -->
```python
assert len(bottinTelephonique) == 6, "Devrait y avoir 6 entrées dans le bottin"
assert bottinTelephonique[0] == '514-555-3333 Eric Fillion',    'Premier numéro invalide'
assert bottinTelephonique[1] == '514-555-2222 Guillaume Huard', 'Deuxième numéro invalide'
assert bottinTelephonique[2] == '418-555-1111 Kim Labrecque',   'Troisième numéro invalide'
assert bottinTelephonique[3] == '418-333-3333 Ian Jacob',       'Quatrième numéro invalide'
assert bottinTelephonique[4] == '418-332-2222 Charles Demers',  'Cinquième numéro invalide'
assert bottinTelephonique[5] == '418-331-1111 Adam Bernard',    'Sixième numéro invalide'
```

---
#### Question #3 (1.5 point)

Créer une fonction `extraireCours(liste,noProgramme)` qui, à partir d'une liste de cours, extrait les cours d'un numéro de programme et trie les cours extraits en ordre croissant de numéro.  La fonction retourne la nouvelle collection.

> La collection reçue en paramètre ne doit pas être modifiée (ni en taille, ni au niveau de l'ordre original).

Résultats attendus:

```
420-2X3-RA Carrières informatiques
420-2Y6-RA Installation d’ordinateurs
420-2Z6-RA Introduction à la programmation
```

<!-- grade_id: q3-code, points: 0.5, grade: true, solution: true -->
```python
def extraireCours(liste,noProgramme):
    listeCours = []
    for unCours in liste:
        if unCours.startswith(str(noProgramme)):
            listeCours.append(unCours)
    listeCours.sort()
    return listeCours

listeCours = [
    "420-2Y6-RA Installation d’ordinateurs",
    "420-2X3-RA Carrières informatiques",
    "420-2Z6-RA Introduction à la programmation",
    "201-273-RA Mathématiques appliquées à l’informatique I",
    "340-101-MQ Philosophie et rationalité",
    "601-101-MQ Écriture et littérature",
    "604-COM-R4 Anglais commun"
]
cours420 = extraireCours(listeCours, 420)
for unCours in cours420:
    print(unCours)
```

Output:
```
420-2X3-RA Carrières informatiques
420-2Y6-RA Installation d’ordinateurs
420-2Z6-RA Introduction à la programmation
```

<!-- grade_id: q3-test, points: 1, grade: true, solution: false -->
```python
assert len(listeCours) == 7, 'liste de cours conserve sept (7) cours'
assert listeCours[0] == "420-2Y6-RA Installation d’ordinateurs", '1er cours -- liste d''origine'
assert listeCours[1] == "420-2X3-RA Carrières informatiques", '2e cours -- liste d''origine'
assert listeCours[2] == "420-2Z6-RA Introduction à la programmation", '3e cours -- liste d''origine'
assert listeCours[3] == "201-273-RA Mathématiques appliquées à l’informatique I", '4e cours -- liste d''origine'
assert listeCours[4] == "340-101-MQ Philosophie et rationalité", '5e cours -- liste d''origine'
assert listeCours[5] == "601-101-MQ Écriture et littérature", '6e cours -- liste d''origine'
assert listeCours[6] == "604-COM-R4 Anglais commun", '7e cours -- liste d''origine'
extr = extraireCours(listeCours, 420)
assert len(extr) == 3, 'Trois cours du programme 420'
assert extr[0] == '420-2X3-RA Carrières informatiques', '1er cours extraction 420'
assert extr[1] == '420-2Y6-RA Installation d’ordinateurs', '2e cours extraction 420'
assert extr[2] == '420-2Z6-RA Introduction à la programmation', '3e cours extraction 420'
extr = extraireCours(listeCours, 601)
assert len(extr) == 1, 'Un cours pour programme 601'
assert extr[0] == '601-101-MQ Écriture et littérature'
extr = extraireCours(listeCours, 123)
assert len(extr) == 0, 'Aucun programme 123'
```

---
#### Question #4 (1.5 point)

À l'aide d'un comparateur externe, trier en ordre croissant de périodes en classe par semaine (somme des deux premiers chiffres) la liste des cours.  

Ex: Carrières informatiques a trois périodes par semaine (théorie: 1 + lab: 2) tandis que le cours de philo est un quatre (4) heures par semaine.

> Aucun traitement particulier à réaliser si deux cours ont le même nombre d'heures-contact par semaine.

Résultats attendus:

```
- Carrières informatiques                       (420-2X3-RA) -- 3 périodes
- Mathématiques appliquées à linformatique I    (201-273-RA) -- 3 périodes
- Anglais commun                                (604-COM-RA) -- 3 périodes
- Philosophie et rationalité                    (340-101-MQ) -- 4 périodes
- Écriture et littérature                       (601-101-MQ) -- 4 périodes
- Installation dordinateurs                     (420-2Y6-RA) -- 6 périodes
- Introduction à la programmation               (420-2Z6-RA) -- 6 périodes
```

<!-- grade_id: q4-code, points: 0.5, grade: true, solution: true -->
```python
def cmp_to_key(mycmp):
    'Convert a cmp= function into a key= function'
    class K(object):
        def __init__(self, obj, *args):
            self.obj = obj
        def __lt__(self, other):
            return mycmp(self.obj, other.obj) < 0
        def __gt__(self, other):
            return mycmp(self.obj, other.obj) > 0
        def __eq__(self, other):
            return mycmp(self.obj, other.obj) == 0
        def __le__(self, other):
            return mycmp(self.obj, other.obj) <= 0  
        def __ge__(self, other):
            return mycmp(self.obj, other.obj) >= 0
        def __ne__(self, other):
            return mycmp(self.obj, other.obj) != 0
    return K

def comparateurDureeCours(coursX, coursY):
    dureeCoursX = coursX[0] + coursX[1]
    dureeCoursY = coursY[0] + coursY[1]
    return dureeCoursX - dureeCoursY
    
cours = [
    [2,4,3,'420-2Y6-RA', 'Installation d''ordinateurs'],
    [1,2,1,'420-2X3-RA', 'Carrières informatiques'],
    [3,3,3,'420-2Z6-RA', 'Introduction à la programmation'],
    [2,1,1,'201-273-RA', 'Mathématiques appliquées à l''informatique I'],
    [3,1,3,'340-101-MQ', 'Philosophie et rationalité'],
    [2,2,3,'601-101-MQ', 'Écriture et littérature'],
    [2,1,3,'604-COM-RA', 'Anglais commun']
]

cours.sort()
for unCours in cours:
    print("- %-45s (%s) -- %i périodes" % (unCours[4], unCours[3], unCours[0] + unCours[1]))
```

Output:
```
- Carrières informatiques                       (420-2X3-RA) -- 3 périodes
- Mathématiques appliquées à linformatique I    (201-273-RA) -- 3 périodes
- Anglais commun                                (604-COM-RA) -- 3 périodes
- Écriture et littérature                       (601-101-MQ) -- 4 périodes
- Installation dordinateurs                     (420-2Y6-RA) -- 6 périodes
- Philosophie et rationalité                    (340-101-MQ) -- 4 périodes
- Introduction à la programmation               (420-2Z6-RA) -- 6 périodes
```

<!-- grade_id: q4-tests, points: 1, grade: true, solution: false -->
```python
asd =  [
    [2,4,3,'420-2Y6-RA', 'Installation d''ordinateurs'],
    [1,2,1,'420-2X3-RA', 'Carrières informatiques'],
    [3,3,3,'420-2Z6-RA', 'Introduction à la programmation'],
    [2,1,1,'201-273-RA', 'Mathématiques appliquées à l''informatique I'],
    [3,1,3,'340-101-MQ', 'Philosophie et rationalité'],
    [2,2,3,'601-101-MQ', 'Écriture et littérature'],
    [2,1,3,'604-COM-RA', 'Anglais commun']
]
asd.sort(key=cmp_to_key(comparateurDureeCours))
assert len(asd) == 7, "Sept cours"
assert asd[0][3] == '420-2X3-RA'
assert asd[1][3] == '201-273-RA'
assert asd[2][3] == '604-COM-RA'
assert asd[3][3] == '340-101-MQ'
assert asd[4][3] == '601-101-MQ'
assert asd[5][3] == '420-2Y6-RA'
assert asd[6][3] == '420-2Z6-RA'
```

```python

```

```python

```