Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Gaston Itaka Kitambala"
COLLABORATORS = ""
```

---

###### <img style='margin-right:0' src="http://dinfo.ca/logoDptInfo.jpg" width=300>

# Laboratoire #7
---

---
#### Question #1 (1 point)
Créez une liste qui contient les noms des pays du G7.  

La liste est référencée par la variable `paysG7`.

| paysG7      |
|-------------|
| États-Unis  |
| Japon       |
| Allemagne   |
| France      | 
| Royaume-Uni | 
| Italie      |
| Canada      |


<!-- grade_id: q1-code, points: 0, grade: false, solution: true -->
```python
paysG7 = ["États-Unis", "Japon", "Allemagne", "France", "Royaume-Uni", "Italie", "Canada"]
```

<!-- grade_id: q1-tests, points: 1, grade: true, solution: false -->
```python
# Validations
assert 'paysG7' in globals(), 'Aucune variable portant le nom `paysG7`'
assert type(paysG7) == list, 'Mauvais type pour la collection'
assert len(paysG7) == 7, 'Il n\'y a pas sept (7) pays dans votre collection'
assert paysG7[0] == 'États-Unis', 'Premier pays mal orthographié'
assert paysG7[1] == 'Japon', '2e pays mal orthographié'
assert paysG7[2] == 'Allemagne', '3e pays mal orthographié'
assert paysG7[3] == 'France', '4e pays mal orthographié'
assert paysG7[4] == 'Royaume-Uni', '5e pays mal orthographié'
assert paysG7[5] == 'Italie', '6e pays mal orthographié'
assert paysG7[6] == 'Canada', '7e pays mal orthographié'
```

## ---
#### Question #2 (2 points)

La Russie, dont le pays était suspendu du groupe économique depuis 2014, est de retour.

Dans une variable `paysG8`, vous devez **copier** la collection provenant de `paysG7` et ensuite y ajouter ce nouveau pays après la France.

Après l'intervention, `paysG7` contiendra sept pays et `paysG8` en contiendra huit.

<!-- grade_id: q2-code, points: 0, grade: false, solution: true -->
```python
# VOTRE CODIFICATION (et supprimer la ligne suivante
paysG8 = paysG7.copy()
paysG8.insert(4, "Russie")
print(paysG7)
print(paysG8)
```

Output:
```
['États-Unis', 'Japon', 'Allemagne', 'France', 'Royaume-Uni', 'Italie', 'Canada']
['États-Unis', 'Japon', 'Allemagne', 'France', 'Russie', 'Royaume-Uni', 'Italie', 'Canada']
```

<!-- grade_id: q2-tests, points: 2, grade: true, solution: false -->
```python
# Validations
assert 'paysG7' in globals(), 'Aucune variable portant le nom `paysG7`'
assert type(paysG7) == list, 'Mauvais type pour la collection'
assert len(paysG7) == 7, 'Il n\'y a pas sept (7) pays dans votre collection'
assert paysG7[3] == 'France', 'Attendu: France'
assert paysG7[4] == 'Royaume-Uni', 'Attendu: Royaume-Uni'
assert paysG8[0] == 'États-Unis', 'Premier pays mal orthographié'
assert paysG8[1] == 'Japon', '2e pays mal orthographié'
assert paysG8[2] == 'Allemagne', '3e pays mal orthographié'
assert paysG8[3] == 'France', '4e pays mal orthographié'
assert paysG8[4] == 'Russie', '5e pays mal orthographié'
assert paysG8[5] == 'Royaume-Uni', '6e pays mal orthographié'
assert paysG8[6] == 'Italie', '7e pays mal orthographié'
assert paysG8[7] == 'Canada', '8e pays mal orthographié'
```

---
#### Question #3 (1 point)

Vous devez afficher, en utilisant une boucle, chaque pays du G7 comme suit:

>    ________ est membre du G7

où le souligné est remplacé par le nom du pays    

<!-- grade_id: cell-7e833c5c32d21094, points: 1, grade: true, solution: true -->
```python
# VOTRE CODIFICATION (et supprimer la ligne suivante
for unPays in paysG7:
    print("%s est membre du G7" % unPays)
```

Output:
```
États-Unis est membre du G7
Japon est membre du G7
Allemagne est membre du G7
France est membre du G7
Royaume-Uni est membre du G7
Italie est membre du G7
Canada est membre du G7
```

## ---
#### Question #4 (2 points)

Créez une fonction `calculerTaxes(listePrix)` qui reçoit un collection de prix et retourne la taxe calculée pour chacun dans une nouvelle collection.  La collection d'origine n'est pas modifiée par votre méthode.

Le calcul doit être arrondi à 2 décimales.  

> Ex: `round(4.567,2)` donne 4.57

Stratégie:
* créer une nouvelle collection vide
* boucle pour visiter chaque prix
* calculer la taxe, faire l'arrondi à 2 décimales et ajouter ce montant à la nouvelle collection
* retourner la nouvelle collection à la fin de la fonction

<!-- grade_id: q4-code, points: 0, grade: false, solution: true -->
```python
def calculerTaxes(listePrix):
    taxes = []
    for unPrix in listePrix:
        taxeCalculee = round(0.14975 * unPrix, 2)
        taxes.append(taxeCalculee)
    return taxes
```

```python
calculerTaxes((10.00, 100.00, 50.00, 40.00, 30.00, 20.00))
# Attendu: [1.5, 14.97, 7.49, 5.99, 4.49, 3.0]
```

Output:
```
[1.5, 14.97, 7.49, 5.99, 4.49, 3.0]
```

<!-- grade_id: q4-tests, points: 2, grade: true, solution: false -->
```python
# Validations
assert 'calculerTaxes' in globals(), 'Aucune fonction portant le nom `calculerTaxes`'
import inspect
assert inspect.isfunction(calculerTaxes), '`calculerTaxes` n\'est pas une fonction'
assert len(inspect.signature(calculerTaxes).parameters) == 1, 'La fonction doit avoir un paramètre'
assert list(inspect.signature(calculerTaxes).parameters.keys())[0] == 'listePrix', 'Le paramètre doit porter le nom `listePrix`'

```

---
#### Question #5 (2 points)

Créez une fonction `denombrerElevesPresents(listePresences)` qui compte le nombre d'items qui sont "présent".

Retourne 0 pour une collection vide ou si tous les élèves sont absents!

<!-- grade_id: q5-code, points: 0, grade: false, solution: true -->
```python
def denombrerElevesPresents(listePresences):
    elevesPresents = listePresences.count("présent")
    return elevesPresents
            
```

```python
print("Attendu: 3, réel: ", denombrerElevesPresents(("présent", "absent", "absent", "présent", "présent")) )
print("Attendu: 5, réel: ", denombrerElevesPresents(("présent", "présent", "présent", "présent", "présent")) )
print("Attendu: 0, réel: ", denombrerElevesPresents(("absent", "absent", "absent", "absent", "absent")) )
print("Attendu: 0, réel: ", denombrerElevesPresents([]) )
```

Output:
```
Attendu: 3, réel:  3
Attendu: 5, réel:  5
Attendu: 0, réel:  0
Attendu: 0, réel:  0
```

<!-- grade_id: q5-tests, points: 2, grade: true, solution: false -->
```python
# Validations
assert denombrerElevesPresents(("présent", "absent", "absent", "présent", "présent")) == 3, 'Trois élèves présents sur cinq'
assert denombrerElevesPresents(("présent", "présent", "présent", "présent", "présent")) == 5, 'Tous les élèves sont pourtant présents'
assert denombrerElevesPresents(("absent", "absent", "absent", "absent", "absent")) == 0, 'Aucun élève présent si tous absents'
assert denombrerElevesPresents([]) == 0, 'Aucun élève présent lorsque la liste est vide'
```

---
#### Question #6 (2 points)

Déclarez une fonction `extraireMinMax(valeurs)` qui retire l'élément ayant la valeur minimale et l'élément ayant la valeur maximale d'une collection.  La collection reçue en paramètre est modifiée (aucun énoncé return).

Stratégie:
* trouver l'élément min
* trouver l'élément max
* retirer l'élément min (pas dans une boucle)
* retirer l'élément max (pas dans une boucle)

**CONSIDÉRER**:
* la collection passée en paramètre est modifiable (list et non tuple)
* Il y a au moins deux éléments dans la collection
* Aucun doublon dans les valeurs; toutes les valeurs sont uniques

<!-- grade_id: q6-code, points: 0, grade: false, solution: true -->
```python
def extraireMinMax(valeurs):
    valeurMax = max(valeurs)
    valeurMin = min(valeurs)
    valeurs.remove(valeurMin)
    valeurs.remove(valeurMax)
    return valeurs
    
```

<!-- grade_id: q6-tests, points: 2, grade: true, solution: false -->
```python
# Validations
w = [8,4,2]
extraireMinMax(w)
assert len(w) == 1, '[4] retourne collection de 1 élément'
x = [1,2,3,4,5]
extraireMinMax(x)
assert len(x) == 3, '[1,2,3,4,5] retourne collection de 3 éléments'
y = [-2,-5,0,9,4,1]
extraireMinMax(y)
assert len(y) == 4, '[-2, -5, 0, 9, 4, 1] retourne collection de 4 éléments'
z = [10,100]
extraireMinMax(z)
assert len(z) == 0, '[10,100] retourne collection de 0 élément'
```

```python

```