Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```javascript
NAME = "Gaston Itaka KITAMBALA"
COLLABORATORS = ""
```

Output:
```
''
```

---

<img style='margin-right:0' src="http://dinfo.ca/logoDptInfo.jpg" width=300>

# Laboratoire #4 (JavaScript)
---

#### Préambule (0.5 point)
Déclarez deux (2) variables nommées: 
    - username (type: chaîne de caractères)
    - age (type: numérique)
en donnant votre propre `username` ainsi que votre `age`.

> P.S.: Vous pouvez mettre `15` ans si vous ne voulez pas dévoiler votre âge réel.

<!-- grade_id: cell-126d9aad32d53cfd, points: 0, grade: false, solution: true -->
```javascript
// VOTRE CODIFICATION (et supprimer la ligne suivante)
var username = "itakita";
var age = 15;
```

<!-- grade_id: preambule-correct, points: 0.5, grade: true, solution: false -->
```javascript
// Validations
var assert = require('assert');
assert(typeof username !== 'undefined', 'Première variable mal définie');
assert(typeof age !== 'undefined', 'Seconde variable mal définie');
assert(typeof username === 'string', 'username n\'est pas une chaîne de caractères');
assert(username.trim().length === 7, 'Le username doit avoir sept (7) lettres');
assert(typeof age === 'number', 'Votre âge doit être de type numérique');
assert(age >= 15 && age <= 60, 'Valeur invalide pour votre âge');
```

---
#### Question #1 (1 point)
Déclarez une fonction nommée `doubler` qui retourne le nombre passé en paramètre (`nb`) multiplié par 2.

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```javascript
// VOTRE CODIFICATION (et supprimer la ligne suivante)
function doubler(nb){
    return nb * 2;
}
```

<!-- grade_id: q1-validation, points: 1, grade: true, solution: false -->
```javascript
var assert = require('assert');
// Validations
assert(typeof doubler !== 'undefined', 'Aucune fonction ne porte le nom `doubler`');
assert(typeof doubler === 'function', '`doubler` n\'est pas une fonction');
assert(doubler(1) === 2, 'Le double de 1 doit donner 2');
assert(doubler(-5) === -10, 'Le double de -5 doit donner -10');
assert(doubler(0) === 0, 'Le double de 0 doit donner 0');
assert(doubler(128) === 256, 'Le double de 128 doit donner 256');
```

---
#### Question #2 (0.5 point)
Affichez le double de `21` en invoquant la fonction `doubler` que vous avez défini.

La réponse attendue est `42`

<!-- grade_id: q2, points: 0.5, grade: true, solution: true -->
```javascript
// VOTRE CODIFICATION (et supprimer la ligne suivante)
console.log(doubler(21));
```

Output:
```
42
```

---
#### Question #3 (1 point)
Invoquez la fonction `doubler` en fournissant la valeur `36` en paramètre, puis stockez le résultat dans la variable `r1`.  

Ensuite, affichez:
> Le double de `36` donne `___`

où le souligné est remplacé par le contenu de la variable `r1`.

<!-- grade_id: q3, points: 0.5, grade: true, solution: true -->
```javascript
// VOTRE CODIFICATION (et supprimer la ligne suivante)
r1 = doubler(36);
console.log(`Le double de 36 donne ${r1}`);
```

Output:
```
Le double de 36 donne 72
```

<!-- grade_id: q3-validation, points: 0.5, grade: true, solution: false -->
```javascript
var assert = require('assert');
// Validations
assert(typeof r1 !== 'undefined', 'La variable `r1` doit être définie');
assert(r1 === 72, 'La variable `r1` doit contenir la bonne réponse (72)');
```

---
#### Question #4 (1 point)
Déclarez la méthode `tripler` qui additionne trois (3) fois le paramètre `valeur` et qui retourne le résultat.  N.B.  Ne pas utiliser l'opérateur de multiplication (astérisque) pour effecter ce calcul.

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```javascript
// VOTRE CODIFICATION (et supprimer la ligne suivante)
function tripler(valeur){
    return valeur + valeur + valeur;
}
    
```

<!-- grade_id: q4-validation, points: 1, grade: true, solution: false -->
```javascript
var assert = require('assert');
// Validations
assert(typeof tripler !== 'undefined', 'Aucune fonction ne porte le nom `tripler`');
assert(typeof tripler === 'function', '`tripler` n\'est pas une fonction');
assert(tripler(1) === 3, 'Le triple de 1 doit donner 3');
assert(tripler(-5) === -15, 'Le triple de -5 doit donner -15');
assert(tripler(0) === 0, 'Le triple de 0 doit donner 0');
assert(tripler(128) === 384, 'Le triple de 128 doit donner 384');
```

---
#### Question #5 (1 point)
Invoquez la fonction `tripler` trois (3) fois avec, successivement, les valeurs 3, 5 et 10.  Affichez chaque réponse sous le format:
> Le triple de `____` donne `____`

en considérant que le premier souligné est remplacé par la valeur à tripler et que le second souligné indique la réponse.

**Résultats attendus:**
> Le triple de 3 donne 9<br>
  Le triple de 5 donne 15<br>
  Le triple de 10 donne 30

<!-- grade_id: q5, points: 1, grade: true, solution: true -->
```javascript
// VOTRE CODIFICATION (et supprimer la ligne suivante)

console.log("Le triple de 3 donne " + tripler(3));
console.log("Le triple de 5 donne " + tripler(5));
console.log("Le triple de 10 donne " + tripler(10));
```

Output:
```
Le triple de 3 donne 9
Le triple de 5 donne 15
Le triple de 10 donne 30
```

---
#### Question #6 (1 point)
Créez la fonction `afficherBieresRestantes(nbBieres)` qui affiche:
> `____` canettes de bière sur un mur.  Retirez-en une, il en restera `____`.

Testez en invoquant cette fonction avec comme valeur 99 bières, 98 bières puis 97 bières.

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```javascript
// VOTRE CODIFICATION (et supprimer la ligne suivante)
function afficherBieresRestantes(nbBieres){
   // console.log(nbBieres + " canettes de bière sur un mur.  Retirez-en une, il est restera " + (nbBieres - 1));
    console.log(`${nbBieres} canettes de bière sur un mur.  Retirez-en une, il est restera ${nbBieres - 1}`);
}  
afficherBieresRestantes(99);
afficherBieresRestantes(98);
afficherBieresRestantes(97);
```

Output:
```
99 canettes de bière sur un mur.  Retirez-en une, il est restera 98
98 canettes de bière sur un mur.  Retirez-en une, il est restera 97
97 canettes de bière sur un mur.  Retirez-en une, il est restera 96
```

<!-- grade_id: q6-validation, points: 1, grade: true, solution: false -->
```javascript
var assert = require('assert');
// Validations
assert(typeof afficherBieresRestantes !== 'undefined', 'Aucune fonction ne porte le nom `afficherBieresRestantes`');
assert(typeof afficherBieresRestantes === 'function', '`afficherBieresRestantes` n\'est pas une fonction');
```

---
#### Question #7 (1 point)
Créez une fonction `convertirFahrenheit` qui reçoit en paramètre une température exprimée en degrés Celsius et qui retourne la valeur calculée en degrés Fahrenheit selon la formule:

> $T_{Fahrenheit} = \frac{9\times T_{Celsius}}{5}+32$

Pour valider si vous interprétez bien la formule, 10˚C équivaut à 50°F.

> Invoquez la fonction (avec la valeur 10) après l'avoir déclarée.

<!-- grade_id: q7, points: 0, grade: false, solution: true -->
```javascript
// VOTRE CODIFICATION (et supprimer la ligne suivante)
function convertirFahrenheit(cTemp){
    return 9 * cTemp/5 + 32;
}

convertirFahrenheit(10);
```

Output:
```
50
```

<!-- grade_id: q7-validation, points: 1, grade: true, solution: false -->
```javascript
var assert = require('assert');
// Validations
assert(typeof convertirFahrenheit !== 'undefined', 'Aucune fonction ne porte le nom `convertirFahrenheit`');
assert(typeof convertirFahrenheit == 'function', '`convertirFahrenheit` n\'est pas une fonction');
assert(Math.abs(convertirFahrenheit(30) - 86.0) < 0.1, '-30°C devrait représenter 86°F');
assert(Math.abs(convertirFahrenheit(0) - 32.0)  < 0.1, '-30°C devrait représenter 32°F');
assert(Math.abs(convertirFahrenheit(10) - 50.0) < 0.1, '-30°C devrait représenter 50°F');
assert(Math.abs(convertirFahrenheit(25) - 77.0) < 0.1, '-30°C devrait représenter 77°F');
```

---
#### Question #8 (1 point)
Affichez les messages suivants, en invoquant la fonction de conversion en degrés Fahrenheit.
> -30°C représente `___`°F<br>
  0°C représente `___`°F<br>
  10°C représente `___`°F<br>
  25°C représente `___`°F

en remplaçant les soulignés par le résultat d'invocation de la fonction.

> *P.S.: Calculez à la main chacun les réponses attendues en Fahrenheit et validez par après si vous arrivez à la bonne valeur*

<!-- grade_id: q8, points: 1, grade: true, solution: true -->
```javascript
// VOTRE CODIFICATION (et supprimer la ligne suivante)
console.log(`${-30} °C représente ${convertirFahrenheit(-30)}`);
console.log(`${0} °C représente ${convertirFahrenheit(0)}`);
console.log(`${10} °C représente ${convertirFahrenheit(10)}`);
console.log(`${25} °C représente ${convertirFahrenheit(25)}`);

```

Output:
```
-30 °C représente -22
0 °C représente 32
10 °C représente 50
25 °C représente 77
```

---
#### Question #9 (1 point)
Créez une fonction nommée `calculerMoyenne` qui additionne les deux nombres reçus en paramètre `(a,b)` avant de diviser la somme en deux.  La fonction retourne la moyenne calculée.

<!-- grade_id: q9, points: 0, grade: false, solution: true -->
```javascript
// VOTRE CODIFICATION (et supprimer la ligne suivante)
function calculerMoyenne(a,b){
    return (a + b) / 2;
}
    
```

<!-- grade_id: q9-validation, points: 1, grade: true, solution: false -->
```javascript
var assert = require('assert');
// Validations
assert(typeof calculerMoyenne !== 'undefined', 'Aucune fonction ne porte le nom `calculerMoyenne`');
assert(typeof calculerMoyenne == 'function', '`calculerMoyenne` n\'est pas une fonction');
assert(Math.abs(calculerMoyenne(1,3) - 2.0) < 0.1, 'La moyenne de 1 et 3 donne 2');
assert(Math.abs(calculerMoyenne(-2,-2) + 2.0) < 0.1, 'La moyenne de 1 et 3 donne 2');
assert(Math.abs(calculerMoyenne(50,100) - 75.0) < 0.1, 'La moyenne de 50 et 100 donne 75');
assert(Math.abs(calculerMoyenne(50,-30) - 10.0) < 0.1, 'La moyenne de 50 et -30 donne 10');
```

---
#### Question #10 (1 point)
Vous devez d'abord définir deux variables (`v1` et `v2`) qui possèdent respectivement les valeurs 15.0 et 26.0.

Ensuite, faites le calcul de la moyenne avec la fonction définie plus tôt et stockez le résultat dans une troisième variable nommée `v3`.

Affichez ensuite le message suivant, en vous servant du contenu des variables:
> La moyenne de `___` et `___` donne `___`

> **Pour avoir une décimale, utilisez la méthode 'toFixed(1)'.**

**Résultat attendu (avec une seule décimale):**
> La moyenne de 15.0 et 26.0 donne 20.5

<!-- grade_id: q10, points: 0.5, grade: true, solution: true -->
```javascript
// VOTRE CODIFICATION (et supprimer la ligne suivante)
var v1 = 15.0;
var v2 = 26.0;
var v3 = calculerMoyenne(v1, v2);
console.log(`La moyenne de ${v1.toFixed(1)} et ${v2.toFixed(1)} donne ${v3.toFixed(1)}`);
```

Output:
```
La moyenne de 15.0 et 26.0 donne 20.5
```

<!-- grade_id: q10-validation, points: 0.5, grade: true, solution: false -->
```javascript
var assert = require('assert');
// Validations
assert(typeof v1 !== 'undefined', 'La variable `v1` doit être définie');
assert(typeof v2 !== 'undefined', 'La variable `v2` doit être définie');
assert(typeof v3 !== 'undefined', 'La variable `v3` doit être définie');
assert(v1 == 15.0, 'La variable `v1` doit contenir la valeur 15.0');
assert(v2 == 26.0, 'La variable `v2` doit contenir la valeur 26.0');
assert(Math.abs(v3 - 20.5) < 0.1, 'La variable `v3` doit contenir le résultat, soit 20.5');
```

```javascript

```