Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```javascript
NAME = ""
COLLABORATORS = ""
```

---

<img style='margin-right:0' src="http://dinfo.ca/logoDptInfo.jpg" width=300>

# Pratique #1 -- Carnet JavaScript
---

#### Problème #2
Voici deux implémentation pour déterminer le nombre d'étoiles à accorder à un exercice au primaire.

L'exercice compte pour 10 points.

Ce tableau résume la quantité d'étoiles à accorder:

| &nbsp;           | &nbsp;    |
|------------------|-----------|
| 10 pts           | 3 étoiles |
| entre 9 et 9.9   | 2 étoiles |
| entre 8 et 8.9   | 1 étoile  |
| en bas de 8      | aucun     |

NOTE: Les deux implémentations présentent des erreurs (identifiez-les et apportez les correctifs).

<!-- grade_id: cell-9dea533226ee5503, points: 0, grade: false, solution: true -->
```javascript
// Première implémentation
function compterEtoiles(note) {
    var nbEtoiles;
    if(note == 10.0)
        nbEtoiles = 3;
    else if(note > 8.9)
        nbEtoiles = 2;
    else if(note >= 8.0)
        nbEtoiles = 1;
    else
        nbEtoiles = 0;
    return nbEtoiles;
   
}
```

<!-- grade_id: cell-414939cfaecaa820, points: 0, grade: false, solution: true -->
```javascript
// Seconde implémentation
function denombrerEtoiles(note) {
    if(note >= 8.0) 
        nbEtoiles = 1;
    else 
        nbEtoiles = 0;
    
    if(note >= 9.0)
        nbEtoiles = 2;
    
    if(note == 10.0) 
        nbEtoiles = 3;

     return nbEtoiles;
}
```

<!-- grade_id: cell-125c2a73135eb2a9, points: 2, grade: true, solution: false -->
```javascript
var assert = require('assert');
assert.equal(compterEtoiles(9.5),2)
assert.equal(compterEtoiles(2.0),0)
assert.equal(compterEtoiles(8.6),1)
assert.equal(compterEtoiles(10.0),3)
assert.equal(compterEtoiles(9.5),2)
assert.equal(compterEtoiles(8.0),1)
assert.equal(compterEtoiles(8.9),1)
assert.equal(compterEtoiles(9.0),2)
```

<!-- grade_id: cell-6a381ce4f2bc3470, points: 2, grade: true, solution: false -->
```javascript
var assert = require('assert');
assert.equal(denombrerEtoiles(9.5),2)
assert.equal(denombrerEtoiles(2.0),0)
assert.equal(denombrerEtoiles(8.6),1)
assert.equal(denombrerEtoiles(10.0),3)
assert.equal(denombrerEtoiles(9.5),2)
assert.equal(denombrerEtoiles(8.0),1)
assert.equal(denombrerEtoiles(8.9),1)
assert.equal(denombrerEtoiles(9.0),2)
```

#### Problème #6

Une carte à jouer est représentée par une chaîne de 2 caractères.

|&nbsp;| &nbsp;                      |
|------|-----------------------------|
| 2C   | 2 de coeur
| 9P   | 9 de pique
| XT   | 10 de trèfle
| VC   | valet de coeur
| DK   | dame de carreau
| RT   | roi de trèfle

Les figures sont: '2', '3', '4', '5', '6', '7', '8', '9', 'X', 'V', 'D', 'R', 'A'

Les sortes sont: 'T' (trèfle), 'C' (coeur), 'P' (pique) et 'K' (carreau)

On veut une fonction `determinerValeurCarte(carte)` où:

- l'as vaut 1
- les figures '2' à '9' valent le chiffre correspondant
- les figures 'X', 'V', 'D' et 'R' valent 10

<!-- grade_id: cell-044b2f4a4ab3a999, points: 0, grade: false, solution: true -->
```javascript
function determinerValeurCarte(carte) {
    let figure =  carte[0];
    let valeur = parseInt(figure);
    if(isNaN(valeur))
        valeur = figure == "A" ? 1 : 10;
        if(figure == "A")
            
            valeur = 1;
        else
            valeur = 10;
    return valeur;
}

/*or
function determinerValeurCarte(carte) {
    let figure =  carte[0];
    let valeur = parseInt(figure);
    if(isNaN(valeur))
        if(figure == "A")
            valeur = 1;
        else
            valeur = 10;
    return valeur;
}
```

<!-- grade_id: cell-6e7f689c541ca653, points: 2, grade: true, solution: false -->
```javascript
var assert = require('assert');
assert.equal(determinerValeurCarte('2P'),2);
assert.equal(determinerValeurCarte('3T'),3);
assert.equal(determinerValeurCarte('4C'),4);
assert.equal(determinerValeurCarte('5K'),5);
assert.equal(determinerValeurCarte('6P'),6);
assert.equal(determinerValeurCarte('7T'),7);
assert.equal(determinerValeurCarte('8C'),8);
assert.equal(determinerValeurCarte('9K'),9);
assert.equal(determinerValeurCarte('XP'),10);
assert.equal(determinerValeurCarte('VT'),10);
assert.equal(determinerValeurCarte('DC'),10);
assert.equal(determinerValeurCarte('RK'),10);
assert.equal(determinerValeurCarte('AP'), 1);
```

```javascript

```